# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import getpass
from sqlalchemy import create_engine
import os
import pymysql

myname = getpass.getuser().split('-')[1]
MYTABLE=myname + '_data'

origin = 'slucor2022-instance-1.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port_number = '3306'
name = 'slucor2022'
pwd = 'SLUcor2022'
db = 'hds5210'

conn = create_engine('mysql+pymysql://' + name + ':' + pwd + '@' + origin + '/' + db)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != '')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

https://opendata.arcgis.com/api/v3/datasets/ed0bc689bde246b18835c7529ba4efb4_0/downloads/data?format=csv&spatialRefId=4326

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
df = pd.read_csv('Electronic_Health_Record_EHR_Incentive_Program_Payments_to_Eligible_Hospitals.csv')
df.to_sql(MYTABLE, conn, if_exists='replace')
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (323,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'])
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Business_ZIP_Code` field and `population`'s `Zip` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
table_name = {}
DF = pd.read_sql_query('SELECT * FROM ' + "population", conn)
for num,table_ in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(table_['Tables_in_hds5210'])
        table_name[name] = pd.read_sql("SELECT * FROM "+name,conn)
population = table_name['population']

answer = pd.merge(table_name[MYTABLE],population,left_on = 'Business_ZIP_Code',right_on = 'Zip',how = 'inner')
maximum = answer[(answer['Population'] ==answer['Population'].max())]

answer = [str([maximum[['NPI']]][0]).split(" ")[14]]

```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])

```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python


data_table = data.pivot_table(values = 'total_payments',index= ['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type'],aggfunc = np.ma.sum,fill_value = 0)

data_table.reset_index()

d_t = pd.merge(data_table.reset_index(),population,left_on = 'Business_ZIP_Code',right_on = 'Zip',how = 'inner')

d_t['total_per_person'] = (d_t["total_payments"])/d_t['Population']

answer = new_test.pivot_table(values = 'total_per_person',index= 'Medicaid_EP_Hospital_Type',aggfunc = np.ma.mean,fill_value = 0)

answer.rename(columns={'total_per_person':'Avg_Pay_per_Capita'},inplace=True)




```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```

```python

```

## Now using Pandas instead of SQL

```python

pd_ = data.pivot_table(values = 'total_payments',index= ['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type'],aggfunc = np.ma.sum,fill_value = 0)

pd_.reset_index()

npd_ = pd.merge(pd_.reset_index(),population,left_on = 'Business_ZIP_Code',right_on = 'Zip',how = 'inner')

npd_['total_per_person'] = (npd_["total_payments"])/npd_['Population']

answer = npd_.pivot_table(values = 'total_per_person',index= 'Medicaid_EP_Hospital_Type',aggfunc = np.ma.mean,fill_value = 0)

answer.rename(columns={'total_per_person':'Avg_Pay_per_Capita'},inplace=True)
```

```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week15_assignment.ipynb
    !git commit -a -m "Submitting the week 15 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

```python

```

```python

```